# Relación de ejercicios 1

### Lee detenidamente las siguientes instrucciones y borra `raise NotImplementedError` en todas las celdas

En cada trozo de código debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) indicadas, que por defecto vienen inicializadas a `None`. Puedes usar variables intermedias siempre que el resultado final del cálculo quede guardado exactamente en la variable que indica el enunciado (debes reemplazar `None` el código necesario, pero nunca cambiar el nombre de las variables ya creadas). 

Después de cada ejercicio verás varias líneas de código ya hechas. Ejecútalas todas (no modifiques su código) y te dirán si tu solución es correcta o no, sin darte pistas de cómo se resuelve (el código de los tests no guarda relación con el código de la solución). Si la solución es correcta, no se mostrará nada, pero si es incorrecta, verás un error indicando cuál es el test que tu solución no pasa. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero si tu código pasa con éxito las líneas que ves, puedes estar bastante seguro de que tu solución es correcta. Asegúrate de que todas las celdas de auto-evaluación indican que tu solución es correcta antes de subir el ejercicio a la plataforma. En caso contrario, no dudes en preguntar las dudas a través del foro -siempre explicando la duda en palabras, sin pegar código fuente-.

Una vez finalizada la actividad, guarda tu fichero, ciérralo con File -> Close and Halt, y después vuélvelo a abrir y ejecútalo completo pinchando en Cell -> Run All Cells, y asegúrate de que no se lanza ningún error. De esta manera comprobarás que no has olvidado nada y que es posible ejecutarlo completo desde 0 y sin errores. No se corregirá ningún fichero que tenga errores de sintaxis y no se pueda, como mínimo, ejecutar completamente. No pasa nada si alguna de las comprobaciones lanza errores por ser incorrecta, pero el código de cada ejercicio no puede tener errores de sintaxis y debe al menos poderse ejecutar.

RECUERDA SUBIR CADA UNO DE LOS FICHEROS .ipynb TAL CUAL (sueltos), SIN COMPRIMIR Y SIN CAMBIARLES EL NOMBRE. Los ficheros subidos deben tener **exactamente el mismo nombre de fichero que tenían cuando los recibiste**. No subas ningún PDF ni ningún fichero ZIP ni nada similar. La plataforma ya los separa automáticamente en carpetas que traen el nombre y apellidos del alumno, por lo que NO es necesario que lo pongas en ninguna parte.

**Ejercicio 1 (2 puntos)**

* Crear en la variable `lista` una **lista** de enteros empezando en -4, no mayores que 150 en pasos de 10 elementos
* Crear en la variable `arr` un ndarray con la función `np.array()` y a partir de él, crear en la variable `mascara` una máscara booleana de los elementos tales que *la raíz cuadrada de su valor absoluto es mayor que 5 pero menor que 10*. **No está permitido utilizar bucles**. No olvides los `import` necesarios. Puedes usar la función `np.abs` y `np.sqrt` que están vectorizadas y se pueden aplicar sobre un array directamente. PISTA: puedes usar dos variables auxiliares `cond_1` y `cond_2` para almacenar las máscaras dos condiciones anteriores (mayor que, y menor que) y unirlas con `&` de numpy en un solo vector booleano.
  * Recuerda que `mascara` debe contener un vector de booleanos. **No debe contener números, sino valores boolenaos True o False**.
* Calcula en la variable `n_cond` el número total de elementos que cumplen la condición. *No está permitido asignar el valor a mano*.
* Utilizando la máscara, reemplazar en la propia variable `arr` los elementos que cumplen la condición por *la raíz cuadrada de su valor absoluto*. **No está permitido utilizar bucles**.
* Crear en la variable `matriz` un array (una matriz) de 3 filas que deben ser el vector `arr` tras el reemplazamiento, el doble de dicho vector, y la mitad del mismo, respectivamente. PISTA: utiliza la misma función que has empleado para crear el array, pasándole una lista formada por los 3 vectores indicados, tendrás la matriz. Comprueba sus dimensiones con `shape` después.

<!-- grade_id: ej1_solucion, points: 0, grade: false, solution: true -->
```python
import numpy as np
lista = None
arr = None
cond_1 = None
cond_2 = None
mascara = None
n_cond = None
matriz = None

lista = list(range(-4, 151, 10))

arr = np.array(lista)

cond_1 = np.sqrt(np.abs(arr)) > 5
cond_2 = np.sqrt(np.abs(arr)) < 10

mascara = cond_1 & cond_2
n_cond = np.count_nonzero(mascara)
arr[mascara] = np.sqrt(np.abs(arr[mascara]))
matriz = np.array([arr, 2 * arr, 0.5 * arr])

```

<!-- grade_id: ej1_test, points: 2, grade: true, solution: false -->
```python
import numpy as np
assert(isinstance(lista, list))
assert(len(lista) == 16)
assert(min(lista) == -4 and max(lista) == 146 and sum(lista) == 1136)   # el contenido de la lista es incorrecto
assert(isinstance(arr, np.ndarray) and arr.size == 16)
assert(isinstance(mascara, np.ndarray) and mascara.size == 16 and mascara.dtype.name == "bool")  # array booleano
assert(n_cond == 8)
assert(isinstance(matriz, np.ndarray) and matriz.shape == (3, 16))
assert(arr.sum() == 706)          # suma de los valores de arr tras el reemplazamiento mediante la máscara
assert(matriz.sum() == 2471)
```

**Ejercicio 2 (2 puntos)**

* (a) **Sin utilizar bucles ni `map`**, crear una función `armonica` que reciba un solo argumento, que será un vector numérico de numpy, y calcule la *media armónica* de los datos de dicho vector, como indica la siguiente fórmula:

  - Dado un vector n-dimensional $v = (v_1, ..., v_n)$, se define la media armónica como 
  
  $$H = \frac{n}{\frac{1}{v_1} + \frac{1}{v_2} + ... + \frac{1}{v_n}}$$  siendo $n$ la dimensión (longitud) del vector
  
Recuerda que el operador `/` está vectorizado para vectores de numpy y es posible dividir un número entre un vector. La función para sumar los elementos de un vector se denomina `sum()` y es un método del array (no necesitas `np.sum(array)` aunque también existe). 

* Utilizando la función `np.apply_along_axis`, invocar a la función `armonica` sobre cada **columna** de la matriz `matriz` definida en el ejercicio anterior, almacenando el resultado en la variable `medias_columnas` que al final contendrá un vector de 3 elementos. PISTA: `np.apply_along_axis(funcion, axis, matriz)` sirve para aplicar una función definida por el usuario, la cual *debe estar pensada para recibir como argumento un vector*, a cada una de las dimensiones de un array. Debemos fijar `axis=0` para aplicarla a cada columna, lo cual resultará en un vector de 16 elementos (tantos como columnas) formado por los resultados de invocar 16 veces a la función (sobre cada una de las 16 columnas de la matriz).

<!-- grade_id: ej2_solucion, points: 0, grade: false, solution: true -->
```python
import numpy as np

def armonica(vector):
    n = len(vector)
    return n / np.sum(1 / vector)

copia_matriz = matriz.copy()
medias_columnas = np.apply_along_axis(armonica, axis=0, arr=copia_matriz)
#medias_columnas = armonica(vector)
```

<!-- grade_id: ej2_test, points: 2, grade: true, solution: false -->
```python
import numpy as np
import inspect
armonica_source = inspect.getsource(armonica)
assert("for" not in armonica_source and "while" not in armonica_source)      # no se permite utilizar bucles en la función
assert(isinstance(medias_columnas, np.ndarray) and medias_columnas.size == 16)
assert(round(medias_columnas.mean(), 2) == 37.82)
```

**Ejercicio 3 (2 puntos)**

Crear una función `duplica_primer_valor` que reciba un argumento llamado `d` que será un diccionario donde todas las claves serán de tipo string y el valor asociado a cada clave será de tipo lista. La función debe duplicar el primer elemento de cada lista, añadiéndolo por el principio de nuevo, y debe devolver un nuevo diccionario con el resultado de dicha modificación (por tanto, el diccionario que le hemos pasado como argumento no debe ser modificado). Asumir que las listas siempre tendrán al menos un elemento. Las nuevas listas tendrán por tanto siempre un elemento más que las originales.
* PISTA: utiliza **listas por comprensión** para crear una nueva lista de pares a partir de la lista formada por los pares del diccionario original. La nueva lista de pares creada debe tener en cada tupla su segundo elemento modificado, como el resultado de concatenar el elemento de la posición 0 con la lista original que ya existía asociada a esa clave. Finalmente, crea un nuevo diccionario con la functión `dict` a partir de dicha lista de pares, y devuélvelo.
* PISTA: recuerda que el operador `+` sirve para concatenar listas, y que puedes construir una lista de un solo elemento con el operador `[elemento]` que luego puedes concatenar con la lista original para obtener una nueva lista.
* Puedes probar la función sobre el diccionario `midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}`.

<!-- grade_id: ej3_solucion, points: 0, grade: false, solution: true -->
```python
def duplica_primer_valor(d):
    nuevo_diccionario = {
        clave: [valor[0]] + valor 
            for clave, valor 
            in d.items()
    }
    
    return nuevo_diccionario

midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}
nuevo_midict = duplica_primer_valor(midict)
```

<!-- grade_id: ej3_test, points: 2, grade: true, solution: false -->
```python
midict = {"edad": [22, 45, 32], "nombre": ["Juan", "Paula", "Elena"]}
res = duplica_primer_valor(midict)
assert(isinstance(res, dict) and len(res.items()) == 2)
assert(res["edad"] == [22, 22, 45, 32] and res["nombre"] == ["Juan", "Juan", "Paula", "Elena"])
```

**Ejercicio 4 (2 puntos)**

Implementa una función `quita_nombres_largos` que reciba dos argumentos: el primero llamado `personas` debe ser una lista de diccionarios, cada uno con datos personales de una persona, entre los cuales se encuentra el campo `"nombre"`. El segundo, llamado `max_long`, debe ser un entero y debe tener valor por defecto 8 si no se especifica dicho argumento. La función debe devolver una nueva lista en la que se han descartado aquellas personas cuyo campo `nombre` tenga longitud mayor que `max_long`. Puedes probar tu función con la lista `personas = [{"nombre": "Hermenegilda", "edad": 15}, {"nombre": "José Luis", "edad": 59}, {"nombre": "Marisa", "edad": 42}, {"edad": 33}]`.
* PISTA: utiliza **listas por comprensión con un if dentro**, para recorrer la lista de personas y construir otra lista en la que sólo estará cada persona si la longitud de su nombre es menor o igual que la longitud máxima indicada como argumento.
* PISTA: tener en cuenta que alguno de los diccionarios puede no contener el campo "nombre" en cuyo caso dicha persona no será incluida en la lista resultante, pero *el código no debe lanzar ningún error* en dicho caso. Recuerda la función `get()` de los diccionarios.

<!-- grade_id: ej4_solucion, points: 0, grade: false, solution: true -->
```python
def quita_nombres_largos(personas, max_long = 8):
    personas_filtradas = [
        persona 
        for persona 
        in personas 
            if "nombre" in persona 
            and len(persona["nombre"]) <= max_long
    ]
    return personas_filtradas
    
personas = [
    {"nombre": "Hermenegilda", "edad": 15},
    {"nombre": "José Luis", "edad": 59},
    {"nombre": "Marisa", "edad": 42},
    {"nombre": "Kris"},
    {"edad": 33}
]

personas_filtradas = quita_nombres_largos(personas, 5)

```

<!-- grade_id: ej4_test, points: 2, grade: true, solution: false -->
```python
personas = [{"nombre": "Hermenegilda", "edad": 15}, {"nombre": "José Luis", "edad": 59}, {"nombre": "Marisa", "edad": 42}, {"edad": 33}]
res_8 = quita_nombres_largos(personas)
assert(isinstance(res_8, list) and len(res_8) == 1)
assert(isinstance(res_8[0], dict))
assert(res_8[0]["nombre"] == "Marisa")

res_10 = quita_nombres_largos(personas, 10)
assert(isinstance(res_10, list) and len(res_10) == 2)
assert(isinstance(res_10[0], dict) and isinstance(res_10[1], dict))
assert(res_10[0]["nombre"] == "José Luis" and res_10[1]["nombre"] == "Marisa")
```

**Ejercicio 5 (2 puntos)**

Implementa una función llamada `reemplaza_minimos` que reciba dos argumentos: 
* Un argumento `m` que será una matriz numérica bidimensional de numpy (array bidimensional)
* Un segundo argumento `repl` que será un número real

La función debe devolver **una nueva matriz** en la que se ha reemplazado en cada columna de la matriz el mínimo de esa columna por el valor del argumento `repl`. Para ello, aplica de nuevo la función `np.apply_along_axis(funcion, axis, matriz)` sobre una copia de la matriz `m`, con argumento `axis=0` (que significa "aplicar a cada columna") donde la función del primer argumento debe ser **una función lambda** de un solo argumento `v` (vector numérico de numpy), que devuelva un nuevo vector resultado de reemplazar en el vector columna original el mínimo de ese vector por `repl`. Dentro de dicha función lambda puedes hacer referencia a la variable `repl` de la función externa sin ningún problema.

Puedes probar tu función sobre la matriz `matriz` definida en el Ejercicio 1, con `repl=-10`.

* PISTA: dentro de la función, crea en primer lugar una variable llamada `copia_m` que contenga una copia de `m` (método `m.copy()`), después invoca a `np_apply_along_axis` y finalmente devuelve `copia_m` como resultado de la función `reemplaza_minimos`.
* PISTA: Python no permite utilizar el operador de asignación dentro de una función lambda, pero puedes usar la función `np.put(v, indice, elemento)` que se comporta exactamente como si fuera `v[indice] = elemento`. Esto modifica directamente el vector pasado como argumento a la función lambda (es decir, cada columna de la matriz), lo cual no es ningún problema si invocas a la función `np_apply_along_axis` sobre una **copia** de la matriz *m* (y no sobre la matriz original que ha recibido `reemplaza_minimos`). Para saber el índice donde se encuentra el mínimo del vector puedes usar `argmin()` aplicada al array.

<!-- grade_id: ej5_solucion, points: 0, grade: false, solution: true -->
```python
import numpy as np

def reemplaza_minimos(m, repl):
    # Hacer una copia de la matriz para no modificar la original
    copia_m = m.copy()
    
    # Función lambda para reemplazar el mínimo en cada columna
    def fun_lambda(columna):
        fila = np.argmin(columna)
        np.put(columna, fila, repl)
        return columna
    
    # Aplicar la función a lo largo de cada columna
    nueva_matriz = np.apply_along_axis(fun_lambda, axis = 0, arr = copia_m)
    
    return nueva_matriz

repl = -10
nueva_matriz = reemplaza_minimos(matriz, repl)

```

<!-- grade_id: ej5_test, points: 2, grade: true, solution: false -->
```python
import inspect
import numpy as np
reemplaza_minimos_source = inspect.getsource(reemplaza_minimos)
m_antes = matriz.copy()
assert("lambda" in reemplaza_minimos_source and       # obligatorio utilizar una función lambda 
       "put(" in reemplaza_minimos_source and         # obligatorio utilizar la función np.put
       "copy()" in reemplaza_minimos_source)          # obligatorio utilizar copy() para no modificar el argumento
nueva_m = reemplaza_minimos(matriz, -10)
assert(nueva_m is not None)
assert((matriz == m_antes).all())                     # la matriz original pasada como argumento no puede ser modificada
assert(nueva_m.sum() == 1964)
```

```python

```

```python

```

```python

```